---
id: 587d78ac367417b2b2512af4
title: Utiliser la propriété flex-direction pour créer une colonne
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cZmWeA4'
forumTopicId: 301109
dashedName: use-the-flex-direction-property-to-make-a-column
---

# --description--

Les deux derniers défis ont utilisé la propriété `flex-direction` avec la valeur `row`. Cette propriété permet également de créer une colonne en empilant verticalement les enfants d'un conteneur flexible.

# --instructions--

Ajoutez la propriété CSS `flex-direction` à l'élément `#box-container`, et donnez-lui la valeur `column`.

# --hints--


L'élément `#box-container` doit avoir une propriété `flex-direction` définie sur `column`.

```js
assert($('#box-container').css('flex-direction') == 'column');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: column;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
