---
id: a105e963526e7de52b219be9
title: Union triée
challengeType: 5
forumTopicId: 16077
dashedName: sorted-union
---

# --description--

Écrivez une fonction qui prend deux tableaux ou plus et renvoie un nouveau tableau de valeurs uniques dans l'ordre des tableaux fournis à l'origine.

En d'autres termes, toutes les valeurs présentes dans tous les tableaux doivent être incluses dans leur ordre d'origine, mais sans doublons dans le tableau final.

Les nombres uniques doivent être triés dans leur ordre d'origine, mais le tableau final ne doit pas être trié dans l'ordre numérique.

Consultez les tests d'assertion pour obtenir des exemples.

# --hints--

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])` devrait retourner `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])` devrait retourner `[1, 2, 3, 5]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])` devrait retourner `[1, 2, 3, 5, 4, 6, 7, 8]`.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])` devrait retourner `[1, 3, 2, 5, 4, 6]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])` devrait retourner `[1, 3, 2, 5, 4]`.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js
function uniteUnique(arr) {
  return arr;
}

uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]);
```

# --solutions--

```js
function uniteUnique(arr) {
  return [].slice.call(arguments).reduce(function(a, b) {
    return [].concat(
      a, 
      b.filter(function(e, currentIndex) {
        return b.indexOf(e) === currentIndex && a.indexOf(e) === -1;
      }));
  }, []);
}
```
