---
id: 587d781b367417b2b2512abb
title: Créer une ligne horizontale à l'aide du hr 
challengeType: 0
videoUrl: 'https://scrimba.com/c/c3bR8t7'
forumTopicId: 301049
dashedName: create-a-horizontal-line-using-the-hr-element
---

# --description--

Vous pouvez utiliser la balise `hr` pour ajouter une ligne horizontale sur la largeur de l'élément qui la contient. Cela peut être utilisé pour définir un changement de sujet ou pour séparer visuellement des groupes de contenu.

# --instructions--

Ajoutez une balise `hr` sous la balise `h4` qui contient le titre de la carte.

**Note :** En HTML, la balise `hr` se ferme d'elle-même et n'a donc pas besoin d'une balise de fermeture distincte.

# --hints--

Votre code doit ajouter une balise `hr` au balisage.

```js
assert($('hr').length == 1);
```

La balise `hr` doit se trouver entre le titre et le paragraphe.

```js
assert(code.match(/<\/h4>\s*?<hr(>|\s*?\/>)\s*?<p>/gi));
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4><s>Google</s>Alphabet</h4>

      <p><em>Google a été fondé par Larry Page et Sergey Brin pendant qu'ils étaient <u>étudiants en doctorat</u> à <strong>l' université de Stanford</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4><s>Google</s>Alphabet</h4>
      <hr>
      <p><em>Google a été fondé par Larry Page et Sergey Brin pendant qu'ils étaient <u>étudiants en doctorat</u> à <strong>l' université de Stanford</strong>.</em></p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
