---
id: 587d781a367417b2b2512ab7
title: Utilisez la balise strong pour mettre le texte en gras
challengeType: 0
videoUrl: 'https://scrimba.com/c/ceJNBSb'
forumTopicId: 301080
dashedName: use-the-strong-tag-to-make-text-bold
---

# --description--

Pour mettre du texte en gras, vous pouvez utiliser la balise `strong`. Celle-ci est souvent utilisée pour attirer l'attention sur un texte et symboliser qu'il est important. Avec la balise `strong`, le navigateur applique le CSS de `font-weight : bold;` à l'élément.

# --instructions--

Placez une balise `strong` autour du texte `Stanford University` à l'intérieur de la balise `p` (n'incluez pas le point).

# --hints--

Votre code doit ajouter une balise `strong` au balisage.

```js
assert($('strong').length == 1);
```

La balise `strong` doit se trouver à l'intérieur de la balise `p`.

```js
assert($('p').children('strong').length == 1);
```

La balise `strong` doit entourer les mots `Stanford University`.

```js
assert(
  $('strong')
    .text()
    .match(/^Stanford University\.?$/gi)
);
```

# --seed--

## --seed-contents--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google a été fondé par Larry Page et Sergey Brin alors qu'ils étaient étudiants en doctorat à la Stanford University</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```

# --solutions--

```html
<style>
  h4 {
    text-align: center;
    height: 25px;
  }
  p {
    text-align: justify;
  }
  .links {
    text-align: left;
    color: black;
  }
  .fullCard {
    width: 245px;
    border: 1px solid #ccc;
    border-radius: 5px;
    margin: 10px 5px;
    padding: 4px;
  }
  .cardContent {
    padding: 10px;
  }
  .cardText {
    margin-bottom: 30px;
  }
</style>
<div class="fullCard">
  <div class="cardContent">
    <div class="cardText">
      <h4>Google</h4>
      <p>Google a été fondé par Larry Page et Sergey Brin alors qu'ils étaient étudiants en doctorat à  la <strong>Stanford University</strong>.</p>
    </div>
    <div class="cardLinks">
      <a href="https://en.wikipedia.org/wiki/Larry_Page" target="_blank" class="links">Larry Page</a><br><br>
      <a href="https://en.wikipedia.org/wiki/Sergey_Brin" target="_blank" class="links">Sergey Brin</a>
    </div>
  </div>
</div>
```
