---
id: cf1111c1c11feddfaeb3bdef
title: Additionner deux nombres avec JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cM2KBAG'
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` est un type de données en JavaScript qui représente les données numériques.

Essayons maintenant d'additionner deux nombres en utilisant JavaScript.

JavaScript utilise le symbole `+` comme opérateur d'addition lorsqu'il est placé entre deux nombres.

**Example:**

```js
const myVar = 5 + 10;
```

`myVar` a maintenant la valeur `15`.

# --instructions--

Changez le `0` pour que la somme soit égale à `20`.

# --hints--

`sum` doit être égal à `20`.

```js
assert(sum === 20);
```

Vous devez utiliser l'opérateur `+`.

```js
assert(/\+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'sum = '+z;})(sum);
```

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
