---
id: cf1111c1c11feddfaeb6bdef
title: Diviser un nombre par un autre avec JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cqkbdAr'
forumTopicId: 17566
dashedName: divide-one-number-by-another-with-javascript
---

# --description--

On peut aussi diviser un nombre par un autre.

JavaScript utilise le symbole `/` pour la division.

**Example**

```js
const myVar = 16 / 2;
```

`myVar` a maintenant la valeur `8`.
# --instructions--

Changez le `0` pour que le `quotient` soit égal à `2`.

# --hints--

La variable `quotient` doit être égale à 2.

```js
assert(quotient === 2);
```

Vous devez utiliser l'opérateur `/`.

```js
assert(/\d+\s*\/\s*\d+/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'quotient = '+z;})(quotient);
```

## --seed-contents--

```js
const quotient = 66 / 0;
```

# --solutions--

```js
const quotient = 66 / 33;
```
