---
id: bd7993c9c69feddfaeb7bdef
title: Multiplier deux nombres décimaux avec JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/ce2GeHq'
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

En JavaScript, vous pouvez également effectuer des calculs avec des nombres décimaux, tout comme avec des nombres entiers.

Multiplions deux nombres décimaux ensemble pour obtenir leur produit.

# --instructions--

Changez le `0.0` pour que le produit soit égal à `5.0`.

# --hints--

La variable `produit` doit être égale à `5.0`.

```js
assert(produit === 5.0);
```

Vous devez utiliser l'opérateur `*`.

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(y){return 'produit = '+y;})(produit);
```

## --seed-contents--

```js
const produit = 2.0 * 0.0;
```

# --solutions--

```js
const produit = 2.0 * 2.5;
```
