---
id: cf1231c1c11feddfaeb5bdef
title: Multiplier deux nombres avec JavaScript
challengeType: 1
videoUrl: 'https://scrimba.com/c/cP3y3Aq'
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

On peut aussi multiplier un nombre par un autre.

JavaScript utilise le symbole `*` pour la multiplication de deux nombres.

**Example**

```js
const myVar = 13 * 13;
```

`myVar` aurait la valeur `169`.

# --instructions--

Changez le `0` pour que le produit soit égal à `80`.

# --hints--

La variable `produit` doit être égale à 80.

```js
assert(produit === 80);
```

Vous devez utiliser l'opérateur `*`.

```js
assert(/\*/.test(code));
```

# --seed--

## --after-user-code--

```js
(function(z){return 'product = '+z;})(produit);
```

## --seed-contents--

```js
const produit = 8 * 0;
```

# --solutions--

```js
const produit = 8 * 10;
```
